rule_files:
    - /rules/arcgis_enterprise.rules.yml

evaluation_interval: 30s

tests:
    - interval: 15s
      input_series:
          - series: 'arcgis_service_state{configured_state="started",folder="Hosted",instance="10.0.1.104:8000",job="arcgis-prometheus",service_name="Pruebafgdb",state="started",type="FeatureServer"}'
            values: '0+0x75'
      alert_rule_test:  
          - eval_time: 2m
            alertname: arcgis_rest_service_is_down
            exp_alerts: []
          - eval_time: 3m
            alertname: arcgis_rest_service_is_down
            exp_alerts:
              - exp_labels:
                  severity: critical
                  job: apo34wv.teide.int
                exp_annotations:
                  summary: "CPU alert for host 'apo34wv.teide.int'"
                  description: "CPU usage for host apo34wv.teide.int is at 100.0% for more than 5 minutes"